Until now, we have implemented manually in raw numpy. The goal of this assignment is to make you familiar with PyTorch and work with Convolutional Neural Networks which are better in Image classification. PyTorch is a deep learning framework and working with it is very similar to python. The major advantages it provides are:

- Work with GPUs
- Has automatic differentiation

In this, assignment we will implement a classifier based on Convolutional Neural Network on SVHN dataset. A CNN is a more suitable type of architecture to process image data. Following is the architecture we will build  

```python
import numpy as np
import torch
import dataset
# if torch.cuda.is_available():
#     print("working on gpu!")
#     device = 'cuda'
# else:
#     print("No gpu! only cpu ;)")
#     device = 'cpu'

```

### Arranging the data

In the following cell we will be loading the data that is already downloaded and made available for you in the servers. 

We implement the following parts in the next cell.
- import the pytorch packages which help us in data handling initially. (done it for you!)
- define the train and test transforms to be applied on the input
- create dataloaders for train and test data
- define the parameters for data handling

```python
import torchvision.transforms as transforms

train_transform = None
test_transform = None

train_loader = None
test_loader = None

train_bs = 32
test_bs = 8

```

### Peek into data

In general, it is always important study the data from every angle to build a good discriminative model. It will help us to notice things and give a good understanding of the requirements of our model. In the following cell we just peek into a random batch of images.

```python
import matplotlib.pyplot as plt


## get a batch of data
images, labels = iter(train_loader)
images  = images.numpy()

fig = plt.figure()

for i in np.arange(1, 13):
    ax = fig.add_subplot(3,4,i, frameon=False)
    img = images[i]
    img = img /2 + 0.5
    img = np.transpose(img, (1,2,0))
    plt.imshow(img)
    ax.set_title(labels[i])
```

Output:
```
TypeError                                 Traceback (most recent call last)
<ipython-input-3-53794587075a> in <module>
      3 
      4 ## get a batch of data
----> 5 images, labels = iter(train_loader)
      6 images  = images.numpy()
      7 

TypeError: 'NoneType' object is not iterable
```

### Architecture

We use the following layers in our model.
- a convolution layer for extracting spatial relationships.
- batchnorm layer for normalizing the weights in the hidden layers
- ReLU activation function for the non-linearity between layers
- dropout for regularization
- and finally fully connected layers in the end

#### Model:

we build the following convolutional neural network architecture on the dataset. 

- convolution layer output_channels-16 kernel_size=5 stride=1 padding-2
- batchnormalization layer
- ReLU activation layer
- maxpool layer kernel_size=2 stride=2
- convolution layer output_channels-32 kernel_size=5 stride=1 padding-2
- batchnormalization layer
- ReLU activation layer
- maxpool layer kernel_size=2 stride=2
- fully connected layer 256
- fully connected layer number_of_classes

### Build the model

- We first define a class called Model.
- In init, we will define all the layers that will be used to build the model
- Forward builds the sequential model taking in Input and returns the Output.



<!-- grade_id: cell-377cf97e5126cc8a, points: 0, grade: false, solution: true -->
```python
class Model(nn.module):
    ## init function is the model initializer.
    ## we define all the layers used in our model. 
    def __init__(self, num_classes=8):
        self.conv1 = nn.Conv2D(3, 16, kernel_size=5, stride=1, pad=2)
        self.bn1 = nn.BatchNorm2d(16)
        ''' 
        REST OF THE MODEL HERE
        '''
        ### BEGIN SOLUTION 
        self.conv1 = nn.Conv2D(3, 16, kernel_size=5, stride=1, pad=2)
        self.bn1 = nn.BatchNorm2d(16)
        self.relu = nn.ReLU()
        self.pool = nn.MaxPool2d(kernel_size=2, stride=2)
        self.conv2 = nn.Conv2D(16, 32, kernel_size=5, stride=1, pad=2)
        self.bn2 = nn.BatchNorm2d(32)
        self.fc1 = nn.Linear(512,256)
        self.fc2 = nn.Linear(256, num_classes)
        ### END SOLUTION 
        
    def forward(x):
        ### BEGIN SOLUTION
        x = self.conv1(x)
        x = self.bn1(x)
        x = self.relu(x)
        x = self.pool(x)
        x = self.conv2(x)
        x = self.bn2(x)
        x = self.relu(x)
        x = self.pool(x)
        x = x.view(x.size(0), -1)
        x = self.fc1(x)
        x = self.fc2(x)
        ### END SOLUTION
        '''
        CODE HERE
        '''
        return x
    
```

Output:
```
NameError                                 Traceback (most recent call last)
<ipython-input-4-604ed81c01c3> in <module>
----> 1 class Model(nn.module):
      2     ## init function is the model initializer.
      3     ## we define all the layers used in our model.
      4     def __init__(self, num_classes=8):
      5         self.conv1 = nn.Conv2D(3, 16, kernel_size=5, stride=1, pad=2)

NameError: name 'nn' is not defined
```

```python
batch_size = 16
optimizer = None
learning_rate = None
epochs = None
number_of_classes = 10
classes = {0 :'T-shirt/top',
                  1 :'Trouser',
                  2 :'Pullover',
                  3 :'Dress',
                  4 :'Coat',
                  5 :'Sandal',
                  6 :'Shirt',
                  7 :'Sneaker',
                  8 :'Bag',
                  9 :'Ankle boot'}
```

<!-- grade_id: cell-cd1cb5622fe024a9, points: 0, grade: true, solution: false -->
```python
## define a instance of the class Model and name it as "model"
model = Model()
test_input = torch.randn(1,32,32,3)
out = model(test_input)
### hidden test case how to write ask raghav
assert()
```

Define a loss criterion, In this assignment we will use cross-entropy loss between the predictions and ground truth to estimate the loss. 
- CrossEntropyLoss - https://pytorch.org/docs/stable/nn.html#crossentropyloss

We also define a optimization strategy to update the weights. In this assignment we use Stochastic Gradient descent with Nesterov momentun from the PyTorch package. 

- SGD - https://pytorch.org/docs/stable/optim.html#algorithms (scroll to optim.SGD)

```python
# define the loss
### BEGIN SOLUTION
loss = nn.CrossEntropyLoss()
### END SOLUTION
loss = None

# optimizer for the model, here we use adam optimizer with a learning rate= and rest default parameters 
### BEGIN SOLUTION
optimizer = 
### END SOLUTION
```

The training loop is setup in the following way:

For every batch in the defined number of epochs

- Move the images and labels to the device
- Extract output by passing input through the model 
- pass the output and ground truth to the loss criterion for batch loss
- clear the gradients 
- backpropagate (compute gradients w.r.t the parameters)
- update the parameters with a single optimization step
- update the training loss for plots

repeat

```python
## training loop 

## Number of epochs the model runs
for epoch in range(epochs):
    # Iterate through the batches in the data
    training_loss = 0.0
    for idx,(images,labels)  in enumerate(train_loader):
        ### BEGIN SOLUTION
        
        ### END SOLUTION
        '''
        CODE HERE
        '''
        
        if idx % 1000 == 0:
            print('Cross Entropy Loss: {:0.4f} Epoch [{}/{}] Iter: {}'.format(loss.item(), epoch, epochs, idx))
        
```

### Testing Loop

In the testing loop we don't update the weights. The trained model is tested for all the samples in test data to compute the accuracy and observe how well the model is generalizing to the unseen data. 

The testing loop is setup in the following way: 

For every batch in the testing data

- Turn off the gradients
- Move the images and labels to the device available
- extract output from the model for the input
- compute the prediction class by choosing the one with maximum probability in the predictions.
- Compare the prediction classes with true classes.
- calculate accuracy
- update test_loss for plots

repeat


```python
## Testing Loop

with torch.no_grad:
    test_loss = 0.0
    correct = 0
    total_samples = 0
    for images, labels in test_loader:
        '''
        YOUR CODE HERE
        '''
        total_samples += label.size(0)
    print("Total Accuracy on the Test set: {} %".format(correct/total*100))
    
```

```python
## Visualize the test samples with predicted output and true output
images, labels = iter(test_loader).next()
images = images.numpy()
images = images.to(device)

out = model(images)
_, preds = torch.max(out, dim=1)

if device == 'cuda':
    preds = np.squeeze(preds.numpy())
else:
    preds = np.squeeze(preds.cpu().numpy())

fig = plt.figure()
for i in np.arange(1, 13):
    ax = fig.add_subplot(4, 3, i)
    imshow(images[i])
    ax.set_title("Pred: {} True: {}".format(classes[preds[i]], classes[labels[i]]), 
                color=('green' if preds[i] == labels[i] else 'red'))
```